# Interact Exercise 3

## Imports

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

```python
from IPython.html.widgets import interact, interactive, fixed
from IPython.display import display
```

# Using interact for animation with data

A [*soliton*](http://en.wikipedia.org/wiki/Soliton) is a constant velocity wave that maintains its shape as it propagates. They arise from non-linear wave equations, such has the [Korteweg–de Vries](http://en.wikipedia.org/wiki/Korteweg%E2%80%93de_Vries_equation) equation, which has the following analytical solution:

$$
\phi(x,t) = \frac{1}{2} c \mathrm{sech}^2 \left[ \frac{\sqrt{c}}{2} \left(x - ct - a \right) \right]
$$

The constant `c` is the velocity and the constant `a` is the initial location of the soliton.

Define `soliton(x, t, c, a)` function that computes the value of the soliton wave for the given arguments. Your function should work when the postion `x` *or* `t` are NumPy arrays, in which case it should return a NumPy array itself.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def soliton(x, t, c, a):
    """Return phi(x, t) for a soliton wave with constants c and a."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: interactex03a, points: 2, grade: true, solution: false -->
```python
assert np.allclose(soliton(np.array([0]),0.0,1.0,0.0), np.array([0.5]))
```

To create an animation of a soliton propagating in time, we are going to precompute the soliton data and store it in a 2d array. To set this up, we create the following variables and arrays:

```python
tmin = 0.0
tmax = 10.0
tpoints = 100
t = np.linspace(tmin, tmax, tpoints)

xmin = 0.0
xmax = 10.0
xpoints = 200
x = np.linspace(xmin, xmax, xpoints)

c = 1.0
a = 0.0
```

Compute a 2d NumPy array called `phi`:

* It should have a dtype of `float`.
* It should have a shape of `(xpoints, tpoints)`.
* `phi[i,j]` should contain the value $\phi(x[i],t[j])$.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interactex03b, points: 4, grade: true, solution: false -->
```python
assert phi.shape==(xpoints, tpoints)
assert phi.ndim==2
assert phi.dtype==np.dtype(float)
assert phi[0,0]==soliton(x[0],t[0],c,a)
```

Write a `plot_soliton_data(i)` function that plots the soliton wave $\phi(x, t[i])$. Customize your plot to make it effective and beautiful.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def plot_soliton_data(i=0):
    """Plot the soliton data at t[i] versus x."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
plot_soliton_data(0)
```

<!-- grade_id: interactex03c, points: 2, grade: true, solution: false -->
```python
assert True # leave this for grading the plot_soliton_data function
```

Use `interact` to animate the `plot_soliton_data` function versus time.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: interactex03d, points: 2, grade: true, solution: false -->
```python
assert True # leave this for grading the interact with plot_soliton_data cell
```