# Matplotlib Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
```

## Contour plots of 2d wavefunctions

The wavefunction of a 2d quantum well is:

$$ \psi_{n_x,n_y}(x,y) = \frac{2}{L}
    \sin{\left( \frac{n_x \pi x}{L} \right)} 
    \sin{\left( \frac{n_y \pi y}{L} \right)} $$

This is a scalar field and $n_x$ and $n_y$ are quantum numbers that measure the level of excitation in the x and y directions. $L$ is the size of the well.
    
Define a function `well2d` that computes this wavefunction for values of `x` and `y` that are NumPy arrays.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def well2d(x, y, nx, ny, L=1.0):
    """Compute the 2d quantum well wave function."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: matplotlibex03a, points: 2, grade: true, solution: false -->
```python
psi = well2d(np.linspace(0,1,10), np.linspace(0,1,10), 1, 1)
assert len(psi)==10
assert psi.shape==(10,)
```

The `contour`, `contourf`, `pcolor` and `pcolormesh` functions of Matplotlib can be used for effective visualizations of 2d scalar fields. Use the Matplotlib documentation to learn how to use these functions along with the `numpy.meshgrid` function to visualize the above wavefunction:

* Use $n_x=3$, $n_y=2$ and $L=0$.
* Use the limits $[0,1]$ for the x and y axis.
* Customize your plot to make it effective and beautiful.
* Use a non-default colormap.
* Add a colorbar to you visualization.

First make a plot using one of the contour functions:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: matplotlibex03b, points: 4, grade: true, solution: false -->
```python
assert True # use this cell for grading the contour plot
```

Next make a visualization using one of the pcolor functions:

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: matplotlibex03c, points: 4, grade: true, solution: false -->
```python
assert True # use this cell for grading the pcolor plot
```